```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz8-functions-loops.ipynb")
```

# ⌛️ Quiz 6 - Using Functions

This quiz will evaluate your mastery of using functions in Python. Functions provide a way to isolate code that you want to use repeatedly, and they allow you to pass in data to the code and get data back out of the code.


## Instructions

You will click on the following link to start your quiz. Do not start the quiz before your lab section. We have checks in place to ensure that you do not start the quiz early. If you do, you will receive a 0 on the quiz. Furthermore, it could be considered an academic integrity violation.

[CLICK HERE TO START YOUR QUIZ IN LAB](https://online.wiseattend.com/student/testLink?c=ENGR131&k=0464) 



## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
# DO NOT MODIFY THIS CELL

from ENGR131_Util_2024 import cell_logger, StudentInfoForm, responses, upsert_to_json_file
# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "quiz8"

StudentInfoForm(**responses)
```

**Question 1: Stress on a Beam**

You are a structural engineer responsible for selecting a beam to support a load. You know the beam is subjected to a uniform load. It is your job to create a function that will calculate the minimum beam design specifications to ensure that the beam will not yield under the expected load.

The stress, $\sigma$, on a beam subjected to a uniform load is given by the equation:

$$ 
\sigma = \frac{M y}{I} \tag{1}
$$

where:
- $M$ is the moment of the force applied to the beam (in Newton-meters, Nm),
- $y$ is the distance from the neutral axis to the point at which the stress is being calculated (in meters, m),
- $I$ is the moment of inertia of the beam's cross-sectional area (in meters to the fourth power, m\(^4\)).

When selecting a beam it is important to be able to include a safety factor in the design. The safety factor is a multiplier that is applied to the calculated stress to ensure that the beam will not fail under the expected load.

If the stress on the beam is sigma, the required design specifications including the safety factor are given by:

$$ \sigma_{\text{yield}} = \text{Safety Factor} \times \sigma  \tag{2}$$

Write Python code to do the following:

* Define a function called `beam_stress` which accepts three input arguments: `M`, `y`, and `I`, and a fourth optional argument called `safety_factor` which defaults to 1.0.
* Define a variable called `sigma` and set it equal to the result of equation (1) above.
* Define a variable called `sigma_yield` and set it equal to the result of the equation (2) above.
* Have your function print the response: `The stress on the beam is: <sigma> N/m^2, and the required beam yield strength is <sigma_yield> N/m^2`, where `<sigma>`, and `<sigma_yield>` are the integer values of the variable for the stress on the beam and yield strength respectively.
* Implement your function so that it returns the required yield stress of the beam for a given input.

Provide an example usage of your function by calling it with the following inputs:
* `M = 1000` Nm
* `y = 0.1` m
* `I = 0.0001` m\(^4\)
* `safety_factor = 1.5`

Save the output of the function call to a variable `required_yield_strength`.

Your code should print the response: 

Note: You can control the precision of the output by using string formatting to round the output to a specific number of decimal places. For example, to round the output to 2 decimal places, you can use the following code:

```python
print(f"{sigma:0.2f}")
```

Your code replaces the prompt:  `...`

```python
# Your function for beam_stress goes here
...

# Example usage:
M = ...
y = ...
I = ...
safety_factor = ...

# Call the function with example inputs save to a variable `required_yield_strength`
...
```

```python
grader.check("q1-stress-on-beam")
```

**Question 2: Detecting the First Signal Peak Above a Threshold 8**

Write a Python function `threshold_peak` that takes a sequence of signal strength measurements (represented as a list of integers) as a variable `signal`. The goal is to identify the first measurement that exceeds a predefined `threshold` value, which indicates the occurrence of a significant signal peak. Once a peak exceeding the threshold is found, the program should return the *index* (position) of this measurement in the dataset and the *value* of the peak, as a tuple then terminate the search.


Requirements:
1. Write a function `threshold_peak` that takes two input arguments: `signal` (a list of integers) and `threshold` (an integer).
   
   1. Use a for loop to iterate through the signal_measurements list.
Within the loop, implement a condition to check if the current measurement exceeds the threshold.
   1. Print the index of the first measurement that exceeds the threshold and the value of this measurement. The print statement should read `First signal peak above threshold found at index <index> with value <value>.`, where <index> and <value> are the index and value of the first measurement that exceeds the threshold.
   2. If a measurement exceeding the threshold is found, use a return statement to exit the loop immediately and return a tuple containing the index of the measurement and the value of the measurement.
   3. If no measurement exceeds the threshold, print a message: `No significant signal peak was detected above the threshold of <threshold>.`, where <threshold> is the value of the threshold.
1. Test your function by calling it with the provided inputs. Save the output of the function call to a variable `peak`.

Hints:
* Use the enumerate() function to get both the index and the value of each item when iterating through the list with a for loop.

```python
# Sample signal measurements
signal_measurements = [10, 12, 7, 5, 10, 18, 21, 4, 3, 2, 17, 20]

# Threshold value
threshold = 20

...
```

```python
grader.check("q2-detecting-threshold-signal")
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/adE3mdGL) add link to navigate to the assignment on GitHub classroom.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q1-stress-on-beam (points: 12)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> assert M == 1000, 'The value of M is not correct.'
>>> assert y == 0.1, 'The value of y is not correct.'
>>> assert I == 0.0001, 'The value of I is not correct.'
>>> assert safety_factor == 1.5, 'The value of the safety factor is not correct.'
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_1')
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_2')
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_3')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_4')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_5')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_6')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q1_7')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q2_1')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q2_2')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q2_3')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q2_4')
>>> drexel_jupyter_logger.variable_logger_csv('0, 4', 'q2_5')
>>> scorer = submit_score()
>>> question_id = 'q1_1'
>>> if M == 1000 and y == 0.1 and (I == 0.0001) and (safety_factor == 1.5):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', question_id)
...     score = 1
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': 1, 'max_score': 1, 'student_response': f'{M}, {y}, {I}, {safety_factor}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Test values are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_2'
>>> assert required_yield_strength == 1500000, 'The required yield stress is not calculated correctly.'
>>> if required_yield_strength == 1500000:
...     drexel_jupyter_logger.variable_logger_csv('1, 1', question_id)
...     score = 1
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 1, 'student_response': f'{required_yield_strength}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Test for stress implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_3'
>>> with patch('builtins.print') as mock_print:
...     assert type(beam_stress(1, 2, 3, 1)) == float, 'The function does not return a numerical value of type float.'
...     if type(beam_stress(1, 2, 3, 1)) == float:
...         drexel_jupyter_logger.variable_logger_csv('1, 1', question_id)
...         score = 1
...     else:
...         score = 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 1, 'student_response': f'{type(beam_stress(1, 2, 3, 1))}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Function correctly returns a numerical value.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_4'
>>> (M, y, I, safety_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     assert beam_stress(M, y, I) == 400000.0, 'yield stress with default safety factor is not correct.'
...     if beam_stress(M, y, I) == 400000.0:
...         score = 2
...         drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
...     else:
...         score = 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{beam_stress(M, y, I)}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: default safety factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_5'
>>> (M, y, I, safety_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     out = beam_stress(M, y, I, safety_factor)
...     assert out == 1600000, 'Safety factor implemented correctly.'
...     if out == 1600000:
...         score = 2
...         drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
...     else:
...         score = 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{out}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: safety factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> import inspect
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_6'
>>> signature = inspect.signature(beam_stress)
>>> parameters = signature.parameters
>>> assert len(parameters) == 4, f'Expected 4 parameters, found {len(parameters)}'
>>> safety_factor_param = parameters.get('safety_factor')
>>> assert safety_factor_param is not None, 'Safety factor parameter missing.'
>>> assert safety_factor_param.default != inspect.Parameter.empty, 'Safety factor does not have a default value, thus not optional.'
>>> if len(parameters) == 4 and safety_factor_param is not None and (safety_factor_param.default != inspect.Parameter.empty):
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{len(parameters)}, {safety_factor_param.default}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Input parameters is is correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_7'
>>> expected_message = 'The stress on the beam is: 400000 N/m^2, and the required beam yield strength is 400000 N/m^2'
>>> with patch('builtins.print') as mock_print:
...     (M, y, I) = (2000, 0.4, 0.002)
...     out = beam_stress(M, y, I)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Input parameters is is correctly defined.


Tests for q2-detecting-threshold-signal (points: 14)
```
>>> import drexel_jupyter_logger
>>> import inspect
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_1'
>>> signature = inspect.signature(threshold_peak)
>>> parameters = signature.parameters
>>> assert len(parameters) == 2, f'Expected 2 parameters, found {len(parameters)}'
>>> assert 'threshold_peak' in locals(), 'Function threshold_peak is not defined.'
>>> assert callable(threshold_peak), 'threshold_peak is not callable.'
>>> if len(parameters) == 2 and 'threshold_peak' in locals() and callable(threshold_peak):
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{len(parameters)}, {callable(threshold_peak)}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Function correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_2'
>>> expected_message = 'First signal peak above threshold found at index 3 with value 22.'
>>> with patch('builtins.print') as mock_print:
...     new_signal = [5, 7, 9, 22, 15, 19]
...     (index, value) = threshold_peak(new_signal, threshold)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Test for stress implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_3'
>>> with patch('builtins.print') as mock_print:
...     new_signal = [5, 7, 9, 22, 15, 19]
...     pk = threshold_peak(new_signal, threshold)
...     assert isinstance(pk, tuple), 'The function does not return a tuple.'
...     assert pk == (3, 22), 'The function does not return the correct index and value.'
>>> if isinstance(pk, tuple) and pk == (3, 22):
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{pk}, {type(pk)}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: Function correctly returns a numerical value.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_4'
>>> expected_message = 'No significant signal peak was detected above the threshold of 20.'
>>> with patch('builtins.print') as mock_print:
...     threshold = 20
...     out = threshold_peak([1, 2, 3], threshold)
...     mock_print.assert_called_once_with(expected_message)
...     assert None == out, 'The function should return None for both index and value when no peak exceeds the threshold.'
>>> if mock_print.call_args[0][0] == expected_message and None == out:
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{mock_print.call_args[0][0]}, {out}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: default safety factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_5'
>>> with patch('builtins.print') as mock_print:
...     assert peak == (6, 21), 'The peak index and value do not match the expected output.'
>>> if peak == (6, 21):
...     score = 4
...     drexel_jupyter_logger.variable_logger_csv('4, 4', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{peak}'}
>>> scorer.add_response(response)
>>> scorer.submit()
Login successful
Data successfully uploaded to the server
responses successfully uploaded to database
```
success_message: peak does not return the correct value.
